### Generate Random array
Generate a numpy array of shape x with random values as its elements.





```python
import numpy as np
```

<!-- grade_id: cell-6c2481fe9b74106f, points: 0, grade: false, solution: true -->
```python
def generate_array(x):
  '''
  input:
    x: tuple defining shape; Eg., (5, 2, 3)
  output:
    y: numpy array of shape x generated using numpy.random library
        and its elements as random values between [0,1).
  '''
  # YOUR CODE HERE
  y =np.random.random(x)
  return y
```

<!-- grade_id: cell-a69e31f80ecccc9d, points: 4, grade: true, solution: false -->
```python
'''test for generate_array'''
def test_generate_array():
  x = (2,4,5,6)
  y = generate_array(x)
  assert y.shape==x
  assert np.all(y<1)
  assert np.all(y>=0)
  print('Test passed', '\U0001F44D')
test_generate_array()
```

Output:
```
Test passed 👍
```

### Shape
Write a python function to calculate shape of a given array.

<!-- grade_id: cell-384a33e1069301ec, points: 0, grade: false, solution: true -->
```python
def return_shape(x):
  '''
  input:
    x: a numpy array of any dimension
  output:
    y: shape of input array as a tuple.
  '''
  # YOUR CODE HERE
  y = np.shape(x)
  return y
```

<!-- grade_id: cell-bfa0a66fc882c4bd, points: 4, grade: true, solution: false -->
```python
'''test for return_shape'''
def test_return_shape():
  x = np.array([1,2])
  y = return_shape(x)
  assert y == (2,)
  print('Test passed', '\U0001F44D')
test_return_shape()
```

Output:
```
Test passed 👍
```

### Mean and STD
Given a data tensor X,  with axis=0 indexing the samples, find mean and standard deviation of the samples .

<!-- grade_id: cell-d9578642617a24f1, points: 0, grade: false, solution: true -->
```python
def CalculateMeanStddev(x):
  '''
  input:
    x: a numpy array of any dimension
  output:
    mean: Numpy array with mean of input samples
    stddev: Numpy array with standard deviation of input samples
  
  HINT: Shape of Input is (samples, feature1, feature2, ...)
  
  Return mean,stddev
  '''
  # YOUR CODE HERE
  mu = np.mean(x,axis =0)
  sd = np.std(x,axis = 0)
  return mu,sd
```

<!-- grade_id: cell-c79e831f7d49c625, points: 4, grade: true, solution: false -->
```python
'''test for CalculateMeanStddev'''
def test_CalculateMeanStddev():
  X = np.array([[1,2,3],
                [5,3,1],
                [3,7,9],
                [3,5,6]])
  mean, stddev = CalculateMeanStddev(X)
  assert np.all(np.isclose(mean, np.array([3.  , 4.25, 4.75])))
  assert np.all(np.isclose(stddev, np.array([1.41421356, 1.92028644, 3.03108891]) ))

  print('Test passed', '\U0001F44D')
test_CalculateMeanStddev()
```

Output:
```
Test passed 👍
```

### Index
Given a matrix of size (m,n) find the element at index position i = a, j =b (start indexing from 0)
(Note: If m<a+1 or n<b+1, then return 'None') 


<!-- grade_id: cell-fe98bacc1889d4a1, points: 0, grade: false, solution: true -->
```python
def element_at_given_index_position(x,a,b):
  '''
  input:
    x: numpy array of shape (m,n)
    a: integer value (row)
    b: integer value (column)
  output:
    y: element at index position (a,b)
  '''
  # YOUR CODE HERE
  dimension = np.shape(x)
  m = dimension[0]
  n = dimension[1]
  if m < a+1 or n< b+1:
    return None
  else:
    return x[a,b]
```

<!-- grade_id: cell-e80a9dacd0a68e56, points: 4, grade: true, solution: false -->
```python
'''test for  element_at_given_index_position'''
def test_element_at_given_index_position():
  
  
  x = np.array([[1,2,3],
                [5,3,1],
                [3,7,9],
                [3,5,6],
                [1,2,3],
                [5,7,8]])
  z =  element_at_given_index_position(x,5,2)
  assert z ==8
  print('Test passed', '\U0001F44D')
  
test_element_at_given_index_position()
```

Output:
```
Test passed 👍
```

### Batch splitting
Given a data matrix X with axis = 0 for samples, split X into batches of 'm' samples each.
Note: last batch can have  less than 'm' samples.
Output should be list of numpy arrays.

<!-- grade_id: cell-59ffcd177dbb9bc9, points: 0, grade: false, solution: true -->
```python
def split_data(x,m):
  '''
  input:
    x: 2d numpy array of any shape (with axis = 0 indicating samples)
    m: integer (sample size iin each batch)
  output:
    y: a list containing numpy arrays with 'm' samples in each array and 
    can be less than m only in last array
  '''
  # YOUR CODE HERE
  z=[]
  x1=[]
  for i in range(0,x.shape[0],m):
    #print(i,m)
    if(i<x.shape[0]):
      x1=x[i:(i+m),:]
      #print((x1))
      z.append(x1)
      #print(z)
    else:
      x2=x[(i-m):(x.shape[0]),:]
      #print(x2)
      z=np.append(z,x2)
  return z
  
  
    
```

<!-- grade_id: cell-7198a438982a4595, points: 4, grade: true, solution: false -->
```python
'''test for  element_at_given_index_position'''
def test_split_data():
  
  
  x = np.array([[1,2,3],
                [5,3,1],
                [3,7,9],
                [3,5,6],
                [11,12,13],
                [5,7,8],
                [1,2,3],
                [5,3,1],
                [3,7,9],
                [3,5,6],
                [11,12,13],
                [5,7,8]])
  z = split_data(x,5)
  assert np.all(z[0]==np.array([[ 1,  2,  3],
       [ 5,  3,  1],
       [ 3,  7,  9],
       [ 3,  5,  6],
       [11, 12, 13]]))
  
  assert np.all(z[1]==np.array([[5, 7, 8],
       [1, 2, 3],
       [5, 3, 1],
       [3, 7, 9],
       [3, 5, 6]]))
  
  assert np.all(z[2]==np.array([[11, 12, 13],
       [ 5,  7,  8]]))

  print('Test passed', '\U0001F44D')
  
test_split_data()
```

Output:
```
Test passed 👍
```

### Shuffle
Given a data tensor X with axis = 0 indexing the sample, randomly shuffle the samples.

<!-- grade_id: cell-1a25ff05d00a08cd, points: 0, grade: false, solution: true -->
```python
def shuffle_samples(x):
  '''
  input:
    x: any general numpy array of dimension (m,n)
  output:
    y: a numpy array with dimension same as of input array but 
    samples are shuffled randomly.
  '''
  # YOUR CODE HERE
  y= np.random.shuffle(x)
  return y
```

<!-- grade_id: cell-2f99f7a0a4672738, points: 4, grade: true, solution: false -->
```python
def test_shuffle_samples():
  x = np.array([[1,2,3],
                [5,3,1],[6,8,0]])
  shuffle_samples(x)
  assert x.shape == (3,3)
  print('Test passed', '\U0001F44D')
  
test_shuffle_samples()
```

Output:
```
Test passed 👍
```

<!-- grade_id: cell-3f341761f142d102, points: 0, grade: false, solution: true -->
```python
def oneHot(y, Ny):
  import numpy as np
  
  '''
  Input:
      y: an int in {0, 1, 2}
      Ny: Number of classes, e.g., 3 here.
  Output:
      Y: a vector of Ny (=3) tuples
    '''
  # YOUR CODE HERE
  Y = np.zeros(Ny)
  Y[y] =1
  return Y
  
```

<!-- grade_id: cell-5b1e3b5f6b6e1740, points: 4, grade: true, solution: false -->
```python
def test_oneHot():
    assert np.all(oneHot(0,3)==np.array([1,0,0]))
    assert np.all(oneHot(1,3)==np.array([0,1,0]))
    assert np.all(oneHot(2,3)==np.array([0,0,1]))
    assert Y_train.shape[1]==3
    print("Test passed', '\U0001F44D")
test_oneHot()
```

Output:
```
<ipython-input-27-3f2f8e248448> in <module>()
      5     assert Y_train.shape[1]==3
      6     print("Test passed', '\U0001F44D")
----> 7 test_oneHot()

<ipython-input-27-3f2f8e248448> in test_oneHot()
      3     assert np.all(oneHot(1,3)==np.array([0,1,0]))
      4     assert np.all(oneHot(2,3)==np.array([0,0,1]))
----> 5     assert Y_train.shape[1]==3
      6     print("Test passed', '\U0001F44D")
      7 test_oneHot()

NameError: name 'Y_train' is not defined
```

### 1-hot encoding
Given three classes, encode them as one hot encoded vectors. For example, if
y = [0, 1, 2] then one hot encoded y = [ [1, 0, 0] , [0, 1, 0] , [0, 0, 1] ]

Given desired output in a list, convert output into one hot encoded vectors.

<!-- grade_id: cell-99c6e0b52fd80b1e, points: 0, grade: false, solution: true -->
```python
def one_hot_encoding_output(x):
  '''
  input:
    x: a list containing desired output, for exmple x = ['cat', 'dog', 'mouse']
  output:
    y: numpy array with one hot encoded vectors as its elements.
  '''
  # YOUR CODE HERE
  dict = {}
  count =0
  ln=[]
  for i in range(len(x)):
    if x[i] not in dict:
      dict[x[i]] =count
      ln.append(count)
      count =count+1
    else:
      a = dict[x[i]]
      ln.append(a)
  #print(ln)
  nclasses = len(set(ln))
  #print(nclasses)
  onehot = []
  for i in range(len(ln)):
    sample =[]
    for j in range(nclasses):
      sample.append(0)
    ind = ln[i]
    sample[ind] =1
    onehot.append(sample)
  hot = np.array(onehot)
  #print(hot)
  return hot

#one_hot_encoding_output(['C','D','M','C','C','M'])

```

<!-- grade_id: cell-1d91cca0bc4c58c0, points: 4, grade: true, solution: false -->
```python
def test_one_hot_encoding_output():
    x = ['cat', 'dog', 'mouse']
    one_hot_encoding_output(x)[0]
    assert np.all(one_hot_encoding_output(x)[2]==np.array([0., 0., 1.]))
    assert np.all(one_hot_encoding_output(x)[1]==np.array([0., 1., 0.]))
    assert np.all(one_hot_encoding_output(x)[0]==np.array([1., 0., 0.]))
    print("Test passed', '\U0001F44D")
test_one_hot_encoding_output()
```

Output:
```
Test passed', '👍
```

Given a data tensor X, with axis = 0 indexing the samples, randomly pick 10 samples and return as a numpy tensor.
If there are less than 10 samples, pad the output tensor with samples having all values as zero(hint. - padded samples will have same dimension as of input samples)


<!-- grade_id: cell-56c3ff271d6e1713, points: 0, grade: false, solution: true -->
```python
def pick_samples(x):
  '''
  input:
    x: a numpy array of shape (m,n) (Note: it may have less than 10 samples also. Your code should be fit for 
    both conditions i.e. m<10 or m>10 )
  output:
    y: numpy array containing 10 samples selected by randomly shuffling input array and padding samples with all 
    elements as zeros if input array has less than 10 samples.
    .
  '''
  # YOUR CODE HERE
  np.random.shuffle(x)
  print(x)
  if (np.size(x,0)>10):
    return x[:10,0:]
  else:
    return np.zeros((10,np.size(x,1)))
```

<!-- grade_id: cell-cfa220cc907612f3, points: 2, grade: true, solution: false -->
```python
def test_pick_sample():
    x = np.array([[1,2,3],
                 [5,3,1],
                 [3,7,9]])
    z = pick_samples(x)
    assert z.shape==(10,x.shape[1])
    print("Test passed', '\U0001F44D")
test_pick_sample()
```

Output:
```
[[1 2 3]
 [3 7 9]
 [5 3 1]]
Test passed', '👍
```

<!-- grade_id: cell-3929fecffdb63091, points: 2, grade: true, solution: false -->
```python
def test_pick_sample():
   
  x = np.array([[1,2,3],
                [5,3,1]])
  Z = np.array
  assert np.all(one_hot_encoding_output(x)[2]==np.array([0., 0., 1.]))
  assert np.all(one_hot_encoding_output(x)[1]==np.array([0., 1., 0.]))
  assert np.all(one_hot_encoding_output(x)[0]==np.array([1., 0., 0.]))
  print("Test passed', '\U0001F44D")
test_one_hot_encoding_output()
  
```

Output:
```
Test passed', '👍
```

### Confusion Matrix
Given the true output values of a classification task is a numpy vector Ytrue
and the estimated output values is  a numpy vector Yest , your function should return
the confusion matrix. Do not use any standard library.

<!-- grade_id: cell-caafc3be9b586839, points: 0, grade: false, solution: true -->
```python
def draw_confusion_matrix(yest,ytrue):
  '''
  input:
    yest: one dimensional numy array consisting of predicted output values(for example  yest = [ 1, 0, 0, 1, 0, 0, 1, 1, 1, 0 ] )
    ytrue: one dimensional numy array consisting of actual output values)( for example, ytrue = [ 1, 1, 0, 1, 0, 0, 1, 0, 0, 0 ] )
  output:
    y: confusion matrix as a numpy array 
    .
  '''
  # YOUR CODE HERE
  tn=0
  tp=0
  fp=0
  fn=0
  for i in range(len(ytrue)):
    if yest[i] == ytrue[i]:
      if yest[i] == 1:
        tp = tp+1
      else:
        tn =tn+1
    else:
      if yest[i] == 0:
        fn = fn+1
      else:
        fp = fp+1
  matrix = np.array([tn,fn,fp,tp]).reshape((2,2))
  #print(matrix)
  return matrix
```

<!-- grade_id: cell-e856a2de88bc457a, points: 4, grade: true, solution: false -->
```python
''' check for draw confusion matrix'''

from sklearn.metrics import confusion_matrix
from sklearn.metrics import roc_curve
import numpy as np


expected = [ 1, 0, 0, 1, 0, 0, 1, 1, 1, 0 ]
predicted = [ 1, 1, 0, 1, 0, 0, 1, 0, 0, 0 ]
results = confusion_matrix(expected, predicted)
assert np.all(draw_confusion_matrix(expected,predicted) == results)
print("Test passed', '\U0001F44D")

```

Output:
```
Test passed', '👍
```

Given the confusion matrix CM = np.array([[40 20 10],
                                                                          [10 30 5],
                                                                          [2 1 50]])
 

1.   What is the accuracy? Save it as 'acc.'.
2.   What is precision? Save it as 'prec' in a numpy array.
3.   What is recall? Save it as 'recall' in a numpy array
4.   What is F-1 score? Save it as 'F_1' in a numpy array





<!-- grade_id: cell-7b1c1beedf29edc2, points: 0, grade: false, solution: true -->
```python
def confusion_matrix_metrics(CM):
  '''
    input:
      CM: confusion matrix
    output:
      accuracy: Accuracy across all classes
      precision: Numpy array with precision for each class
      recall: Numpy array with recall for each class
      F_1 : Numpy array with f_1 score for each class
      
      Return the results as a tuple (accuracy, precision, recall, f_1)
  '''
  # YOUR CODE HERE
  m,n = np.shape(CM)
  nclasses =m
  total = 0
  for i in range(m):
    for j in range(m):
      total = total+CM[i][j]
  accvar =0
  for i in range(m):
    accvar = accvar+CM[i][i]
  accuracy = accvar/total
  prec = np.empty(m)
  recall = np.empty(m)
  f1 = np.empty(m)
  rowwise = np.sum(CM,axis=1)
  colwise = np.sum(CM,axis=0)
  for i in range(m):
    recall[i] = CM[i][i]/(rowwise[i])
    prec[i] = CM[i][i]/(colwise[i])
    f1[i] = 2*recall[i]*prec[i]/(recall[i]+prec[i])
  return accuracy,prec,recall,f1
  
```

<!-- grade_id: cell-bdf1a16a48bdf3a7, points: 4, grade: true, solution: false -->
```python
''' check for confusion matrix'''

def test_confusion_matrix_metrics():
  CM = np.array([[40,20,10], [10,30,5], [2,1,50]])
  accuracy, precision, recall, F_1 = confusion_matrix_metrics(CM) 
  
  assert np.isclose(accuracy, 0.7142857142857143)
  assert np.all(np.isclose(precision, np.array([0.7692307692307693, 0.5882352941176471, 0.7692307692307693])))
  assert np.all(np.isclose(recall, np.array([0.5714285714285714, 0.6666666666666666, 0.9433962264150944])))
  assert np.all(np.isclose(F_1, np.array([0.6557377049180327, 0.625, 0.847457627118644])))
  print("Test passed', '\U0001F44D")
  
test_confusion_matrix_metrics()
```

Output:
```
Test passed', '👍
```

```python

```